# Artificial Intelligence for Robotics 03

## General Information:
Please do not add or delete any cells. Answers belong into the corresponding cells (below the question). If a function is given (either as a signature or a full function), you should not change the name, arguments or return value of the function.

If you encounter empty cells underneath the answer that can not be edited, please ignore them, they are for testing purposes.

When editing an assignment there can be the case that there are variables in the kernel. To make sure your assignment works, please restart the kernel and run all cells before submitting (e.g. via Kernel -> Restart & Run All).

Code cells where you are supposed to give your answer often include the line `raise NotImplementedError`. This makes it easier to automatically grade answers. If you edit the cell please outcomment or delete this line.

## Submission:
Please submit your notebook via the web interface (in the main view -> Assignments -> Submit). The assignments are **due on Monday at 20:00**.

## Group Work:
Please enter your UID (your username here) and those of your group partners into the next cell. We apply plagiarism checking, so do not submit others solutions! If an assignment has a copied solution, the task will be graded with 0 points for all people with the same solution.

## Questions about the Assignment:
If you have questions about the assignment please post them in the LEA forum before the deadline. Don't wait until the last day to post questions!

### Please add the usernames of all your team members in the manner member1, member2 in next cell (example given below)

member1 = 'example'

member2 = 'example2'

#### If you are not working in a group, then please add member2 as none2s

<!-- grade_id: cell-a9c2f4587c671f98, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
member1 = 'ahoosh2s'
member2 = 'anuhel2s'
#raise NotImplementedError()
```

<!-- grade_id: cell-e27ced94e387fe18, points: 0, grade: true, solution: false -->
```python
# Execute this cell to make sure you correctly filled in the usernames of the team members

def group_name_test():
    for member_id in [member1, member2]:
        assert isinstance(member_id, str), "Please give your member id as a string."
        assert len(member_id) > 0, "You need to fill in the member id for both members"
        assert member_id.endswith("2s"), "The member id should end with 2s (Your JupyterHub username)"

group_name_test() 
print("All tests passed!")
```

Output:
```
All tests passed!
```

# Task 1

**[36 Point(s)]**


# Introduction to Algorithms

You can get help from the excerpt of *Cormen’s* very nice book *Introduction to Algorithms*, available as a PDF ﬁle (`cormen_introalgorithms`) in LEA, to solve the following tasks.

## Task 1.1

**[10 Point(s)]**

### Task 1 [10 Points]

Why is a special notation needed to classify algorithms? Doesnt it suffice to merely measure the runtime in seconds? **Explain**.

<!-- grade_id: algorithms_introduction_A, points: 10, grade: true, solution: true -->
The runtime model is not a comprehensive explanation because it is influenced by software and hardware components. Additionally, runtime does not explain the relationship between the size of the input and the growth of runtime. For example, O(n) is a linear notation, while O(n^m) is exponential in nature, with its growth far exceeding O(n) for larger inputs. The difference between algorithms in terms of runtime cannot be fully explained through runtime results alone, so comparing algorithms in terms of time-efficiency and space-efficiency is impossible using, for example, a table of different runtime tests.

## Task 1.2

**[6 Point(s)]**

### Task 2 [6 Points]

Let $g : \mathbb{N} \rightarrow \mathbb{R}+$ be an arbitrary function. The set of functions $f : \mathbb{N} \rightarrow \mathbb{R}+$, which do not gorw faster than the function $g$ after a specific point $n_0$, is denoted as $O(g(n))$. More specifically:

$$
O(g(n) := f(n) \vert \exists c \in \mathbb{R}, \exists n_o \in \mathbb{N} : 0 \leq f(n) \leq c g(n) \forall n \geq n_0)
$$

**Prove** the following:
- $f(n) = 100n^2 \in O(n^2)$
- $f(n) = n^6 + 100n^5 \in O(n^6)$

<!-- grade_id: algorithms_introduction_B, points: 6, grade: true, solution: true -->
YOUR ANSWER HERE

## Task 1.3

**[3 Point(s)]**

### Task 3 [3 Points]

What is the asymptotic time complexity of the following python code? Assume that `ANY_CONST` is an arbitrary constant in your program, which has assigned a 2D array as its value.

```python
sum_ = 0
for i in range(0, J):
    for j in range(0, K):
        if arr[i][j] <= ANY_CONST:
            sum_ = sum_ + arr[i][j]
            print(sum_)
```

<!-- grade_id: algorithms_introduction_C, points: 3, grade: true, solution: true -->


## Task 1.4

**[15 Point(s)]**

### Task 4 [12 Points]

For each function $f(n)$ and time $t$ in the following table, determine the largest size $n$ of a problem that can be solved in time $t$, assuming that the algorithm to solve the problem takes $f(n)$ microseconds.

Show your **path to the solution** for:
- $log(n)$
- $2^n$
- $nlog(n)$.

| | 1 second | 1 minute | 1 hour | 1 day | 1 month | 1 year | 1 century |
|-|----------|----------|--------|-------|---------|--------|-------------|
| $\log_{10} n$ | | | | | | | |
| $\sqrt{n}$ | | | | | | | |
| $n$ | | | | | | | |
| $n \log_{10} n$ | | | | | | | |
| $n^2$ | | | | | | | |
| $n^3$ | | | | | | | |
| $2^n$ | | | | | | | |
| $n!$ | | | | | | | |

 > If you need to write any code, please do so in the code cell below.

<!-- grade_id: algorithms_introduction_D, points: 15, grade: true, solution: true -->
t=log(n)⟹n=10t

t=2**n ⟹ n=log2(t)

n*log(n) => is calculated numerically


| | 1 second | 1 minute | 1 hour | 1 day | 1 month | 1 year | 1 century |
|-|----------|----------|--------|-------|---------|--------|-------------|
| $\log_{10} n$ |107309 |181445 |1754047 |12038693 |51523268 |138503720 |428796745 |
| $\sqrt{n}$ | |2440 |60000 |1000000 |16000000 |55000000 |179000000 |
| $n$ |1000000 |60000000 |3600000000 |86400000000 |2592000000000 |31536000000000 |3156000000000000 |
| $n \log_{10} n$ |16404 |162888 |1572184 |10000000 |21858983 |61963474 |185000000 |
| $n^2$ |1000 |2449 |60000|9289 |161602 |179993 |1771012 |
| $n^3$ |100 |394 |1442 |1987 |2954 |3110 |14621 |
| $2^n$ |19 |25 |31 |36 |41 |43 |46 |
| $n!$ |12 |14 |15 |16 |20 |22 |23 |



## Task 1.5

**[2 Point(s)]**

<!-- grade_id: algorithms_introduction_E, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
import numpy as np
time_in_microseconds = {
    "1 second": 1_000_000,
    "1 minute": 60_000_000,
    "1 hour": 3_600_000_000,
    "1 day": 86_400_000_000,
    "1 month": 2_592_000_000_000,
    "1 year": 31_536_000_000_000,
    "1 century": 3_156_000_000_000_000,
}

def log_n():
    results = {}
    for item in time_in_microseconds:
        t = time_in_microseconds[item]
        n = 1  # Start from 1 to avoid log(0)
        while n * np.log(n) <= t:  # Check for log(n) <= t
            n += 1
        results[item] = n - 1  # Store the largest n that satisfies log(n) <= t
    return results

# Print results
print("The results for log(n): {}".format(log_n()))

def sqrt_n():
    results = {}
    for item in time_in_microseconds:
        t = time_in_microseconds[item]
        n = 1  # Start from 1 to avoid sqrt(0)
        while np.sqrt(n) <= t:  # Check for sqrt(n) <= t
            n += 1
        results[item] = n - 1  # Store the largest n that satisfies sqrt(n) <= t
    return results

# Print formatted results
print("The results for sqrt(n): {}".format(sqrt_n()))

def calculate_n():
    results = {}
    for item in time_in_microseconds:
        t = time_in_microseconds[item]
        n = t  # The largest n that can be handled in time t
        results[item] = n  # Store n directly, since n can equal t
    return results

# Print formatted results
print("The results for n: {}".format(calculate_n()))

def n_log_n():
    results = {}
    for item in time_in_microseconds:
        t = time_in_microseconds[item]
        n = 1  # Start from 1 to avoid log(0)
        while n * np.log(n) <= t:  # Check for n * log(n) <= t
            n += 1
        results[item] = n - 1  # Store the largest n that satisfies n * log(n) <= t
    return results

# Print formatted results
print("The results for n * log(n): {}".format(n_log_n()))

def n_squared():
    results = {}
    for item in time_in_microseconds:
        t = time_in_microseconds[item]
        n = 1  # Start from 1
        while n**2 <= t:  # Check for n^2 <= t
            n += 1
        results[item] = n - 1  # Store the largest n that satisfies n^2 <= t
    return results

# Print formatted results
print("The results for n^2: {}".format(n_squared()))

def n_cubed():
    results = {}
    for item in time_in_microseconds:
        t = time_in_microseconds[item]
        n = 1  # Start from 1
        while n**3 <= t:  # Check for n^3 <= t
            n += 1
        results[item] = n - 1  # Store the largest n that satisfies n^3 <= t
    return results

# Print formatted results
print("The results for n^3: {}".format(n_cubed()))

def power_of_two():
    results = {}
    for item in time_in_microseconds:
        t = time_in_microseconds[item]
        n = 0  # Start from 0, since 2^0 = 1
        while 2**n <= t:  # Check for 2^n <= t
            n += 1
        results[item] = n - 1  # Store the largest n that satisfies 2^n <= t
    return results

# Print formatted results
print("The results for 2^n: {}".format(power_of_two()))

def factorial_limit():
    results = {}
    for item in time_in_microseconds:
        t = time_in_microseconds[item]
        n = 0  # Start from 0 since 0! = 1
        factorial = 1  # Initialize factorial(0) = 1
        while factorial <= t:  # Check for n! <= t
            n += 1
            factorial *= n  # Calculate n! incrementally
        results[item] = n - 1  # Store the largest n that satisfies n! <= t
    return results

# Print formatted results
print("The results for n!: {}".format(factorial_limit()))

#raise NotImplementedError()
```

# Task 2

**[80 Point(s)]**


# Uninformed Search

In this Assignment you will implement the unfinformed search strategies *Breadth First Search*, *Depth First Search* and *Iterative Deepening Depth First Search*.

## Task 2.1

**[10 Point(s)]**

## Grid World [10 Points]

You will find three text files, each representing a map for the programming assignment in the maps folder. A search agent is supposed to explore each of these three maps using *breadth-first search*, *depth-first search* and *iterative-deepening depth-first search*.

### Grid World
You will find three text files in `./Complexity_and_Uninformed_Search_files/data`, each representing a map. They are to be interpreted as follows:
- The text file is considered a grid.
- Each character in the text file represents a `cell` in this grid.
  - `*`: goal
  - `space`: free
  - `S`: initial position of the agent
  - Any other character represents an obstacle

**Implement a parsing mechanism to load the `.txt` files into a meaningful data structure into memory. Please also fill the docstring of the function and add typehints such that I can understand what arguments are expected and what it returns.**

<!-- grade_id: uninformed_search_A, points: 0, grade: false, solution: true -->
```python
def maze_map_to_tree(maze_map):
    """Function to create a tree from the map file. The idea is
    to check for the possible movements from each position on the
    map and encode it in a meaningful data structure.

    Parameters
    ----------
    maze_map : [String]
        [Takes the file path as a parametre]

    Returns
    -------
    [2D list storing the map and another list conatin the position of the start point]
        [The function returns a 2D list containing all the characters in the map]
    """
    maze = []
    with open(maze_map, 'r') as f:
        for line in f:
            maze.append(list(line.strip()))

    # Find the start node position
    for i in range(len(maze)):
        for j in range(len(maze[i])):
            if maze[i][j] == 's':
                return maze, (i, j)  # Return maze and start node coordinates
    return maze, None
    # YOUR CODE HERE
    #raise NotImplementedError()

# YOUR CODE HERE
#raise NotImplementedError()

```

<!-- grade_id: test_uninformed_search_A0, points: 10, grade: true, solution: false -->
```python
# Test

# Please write your code test here!

```

## Task 2.2

**[10 Point(s)]**

## Agent
The agent must follow the following rules:
- The agent can move from one call to another at each step.
- The agent can **only move to the left, right, up or down** from the current position. The surrounding
cells should be considered as the children of the current cell.
- The agent does **not have any previous knowledge** about the environment (such as goal positions or obstacles) so it has to **explore**.
- The agent **cannot move through obstacles** and the **map is closed**.

## Path Visualization [10 Points]

The function `assign_character` is used to indicate the direction of traversal of the search agent through the nodes of the search tree. Please use the following symbols for visualizing the respective movement of the agent:
 - `U+2500`: ─
 - `U+2502`: │
 - `U+250C`: ┌
 - `U+2510`: ┐
 - `U+2514`: └
 - `U+2518`: ┘
 - `U+251C`: ├
 - `U+2524`: ┤
 - `U+252C`: ┬
 - `U+2534`: ┴
 - `U+253C`: ┼
 - `U+2574`: ╴
 - `U+2575`: ╵
 - `U+2576`: ╶
 - `U+2577`: ╷

With the `write_to_file` function you should be able to save the individual paths to each goal, as well as the final fully explored map with the search path to a `.txt` file.

1. **Implement the `assign_character` function and complete the docstring, such that one can tell how the function is called, what it returns and how it works.**
1. **Implement the `write_to_file` function and comeplete the docstring.**

<!-- grade_id: uninformed_search_B, points: 0, grade: false, solution: true -->
```python
def assign_character(maze_map, current_node, prev_node):
    """Function to assign character for the visited nodes.
    
    Parameters
    ----------
    maze_map : []
        [The list which contains all the characters fetched from the file]
    current_node : [Object of the class node]
        [The current node]
    prev_node : [Object of the class node]
        [The previous node]
    
    Returns
    -------
    it is just replaces the characters in the maze map
    """
    """
    Function to assign a character for visited nodes along the path from start to goal.
    """
    if prev_node is None:
        return # No previous node for the start position

    move_x = current_node.x - prev_node.x
    move_y = current_node.y - prev_node.y

    # Determine the character based on movement direction
    if move_x == 1 and move_y == 0:     # Downward move
        maze_map[prev_node.x][prev_node.y] = '│'
    elif move_x == -1 and move_y == 0:   # Upward move
        maze_map[prev_node.x][prev_node.y] = '│'
    elif move_x == 0 and move_y == 1:    # Rightward move
        maze_map[prev_node.x][prev_node.y] = '─'
    elif move_x == 0 and move_y == -1:   # Leftward move
        maze_map[prev_node.x][prev_node.y] = '─'
    # YOUR CODE HERE
    #raise NotImplementedError()

def path_to_goal(node, maze_map):
    """
    Backtrack from the goal node to the start node and assign path characters.
    """
    while node.parent:  # Until we reach the start node
        assign_character(maze_map, node, node.parent)
        node = node.parent  # Move backward along the path
    maze_map[node.x][node.y] = 's'  # Mark the start position explicitly if needed

def print_path (maze):
    delimiter = ""
    for item in maze:
        str_item = delimiter.join(item)
        print (str_item, '\n')

def write_to_file(file_name, path):
    """Function to write traversed path to each goal, as well as the fully explored map to a `.txt` file.
    
    Parameters
    ----------
    filen_name : string
        This parameter defines the name of the txt file.
    path : String
        the path to the file
    """
    delimiter = ""
    with open(path+file_name, "w", encoding="utf-8") as file:  # Specify UTF-8 encoding
        for item in file:
            str_item = delimiter.join(item)
            file.write(str_item + '\n')
    # YOUR CODE HERE
    #raise NotImplementedError()
```

<!-- grade_id: test_uninformed_search_B0, points: 10, grade: true, solution: false -->
```python
# Test

# Please write your code test here!

```

## Task 2.3

**[20 Point(s)]**

## Breadth First Search [20 Points]

Implement *Breadth-First Search*. Make sure to fill the docstring of the function and add typehints such that one can understand what arguments are expected and what it returns.

Run the search algorithm for each map and save the results into `./data/results/<search_algo_name>_<map_number>.txt`.

<!-- grade_id: uninformed_search_C, points: 0, grade: false, solution: true -->
```python
from collections import deque
class Node:
    def __init__(self, x, y, state, parent=None):
        self.x = x
        self.y = y
        self.state = state
        self.parent = parent  # Set parent to trace the path

def breadth_first_search(maze_map, monitor=False):
    """ This function implements the breadth-first search (BFS) algorithm.
    The key point of the BFS algorithm is the use of a queue like first-in-first-out (FIFO) memory as the frontier.
    Parameters
    ----------
    maze_map : File path of the txt map
        The maze map as it was loaded from the file.
    monitor : bool
        Flag for enabling or disabling command line output for the search run. Defaults to False: No monitoring.
    Returns
    it both returns and prints the modified map list to show the navigation through the map
    The map of the traversed goals in the form of a list
        A list of the nodes containing path to the goal.
    """
    
    movements = [(-1, 0), (1, 0), (0, -1), (0, 1)]
    maze, start_node = maze_map_to_tree(maze_map)
    
    if not start_node:
        print("No start node found in the maze.")
        return

    rows, columns = len(maze), len(maze[0])
    queue = deque([Node(start_node[0], start_node[1], maze[start_node[0]][start_node[1]])])
    visited = set()

    while queue:
        node = queue.popleft()
        
        # Skip if already visited
        if (node.x, node.y) in visited:
            continue
        visited.add((node.x, node.y))

        # Check if goal is found
        if node.state == '*':
            path_to_goal(node, maze)  # Mark path from start to this goal
            maze[node.x][node.y] = ' '  # Replace goal with space after marking path
            continue  # Continue to find additional goals if any

        # Explore neighbors
        for move_x, move_y in movements:
            new_position_x, new_position_y = node.x + move_x, node.y + move_y
            if 0 <= new_position_x < rows and 0 <= new_position_y < columns and (new_position_x, new_position_y) not in visited:
                if maze[new_position_x][new_position_y] in (' ', '*'):
                    # Append neighbor node with current node as parent
                    queue.append(Node(new_position_x, new_position_y, maze[new_position_x][new_position_y], node))
                    
    print_path(maze)
    print("No more goals were found.")
    return maze

breadth_first_search("./Complexity_and_Uninformed_Search_files/data/map1.txt")
#write_to_file("BFS_map1", "./Complexity_and_Uninformed_Search_files/data/")
breadth_first_search("./Complexity_and_Uninformed_Search_files/data/map2.txt")
#write_to_file("BFS_map2", "./Complexity_and_Uninformed_Search_files/data/")
breadth_first_search("./Complexity_and_Uninformed_Search_files/data/map3.txt")
#write_to_file("BFS_map3", "./Complexity_and_Uninformed_Search_files/data/")

    # YOUR CODE HERE
    #raise NotImplementedError()
```

Output:
```
========================================================================== 

|  s                              |                                      | 

|  ──────────────────│            |                                      | 

|====================│            |                                      | 

|                    │                                                   | 

|                 │───────────────────│───────────────────────           | 

|                 │===================│                                  | 

|                 │                   │                                  | 

|                 │                   │                                  | 

|                 │                   ──────────────────────────│        | 

|                 │                   ==========================│        | 

|                 │                                             │        | 

|                 │                                             │        | 

```

<!-- grade_id: test_uninformed_search_C0, points: 20, grade: true, solution: false -->
```python
# Test

# Please write your code test here!

```

## Task 2.4

**[20 Point(s)]**

## Depth First Search [20 Points]

Implement *Depth-First Search*. Make sure to fill the docstring of the function and add typehints such that one can understand what arguments are expected and what it returns.

Run the search algorithm for each map and save the results into `./data/results/<search_algo_name>_<map_number>.txt`.

<!-- grade_id: uninformed_search_D, points: 0, grade: false, solution: true -->
```python
def depth_first_search(maze_map, monitor=False):
    """ This function implements the depth-first search (DFS) algorithm.
    The key point of the DFS algorithm is the use of a stack like last-in-first-out (LIFO) memory as the
    frontier.
    
    Parameters
    ----------
    maze_map : [type]
        [description]
    monitor : bool
        Flag for enabling or disabling command line output for the search run. Defaults to False: No
        monitoring.
    
    Returns
    -------
    [type]
        [description]
    """

    movements = [(-1, 0), (1, 0), (0, -1), (0, 1)]
    maze, start_node = maze_map_to_tree(maze_map)
    
    if not start_node:
        print("No start node found in the maze.")
        return

    rows, columns = len(maze), len(maze[0])
    stack = [Node(start_node[0], start_node[1], maze[start_node[0]][start_node[1]])]
    visited = set()

    while stack:
        node = stack.pop()
        
        # Skip if already visited
        if (node.x, node.y) in visited:
            continue
        visited.add((node.x, node.y))

        # Check if goal is found
        if node.state == '*':
            path_to_goal(node, maze)  # Print path to goal
            maze[node.x][node.y] = ' '  # Replace goal with space
            continue  # Continue to find additional goals if any

        # Explore neighbors
        for move_x, move_y in movements:
            new_position_x, new_position_y = node.x + move_x, node.y + move_y
            if 0 <= new_position_x < rows and 0 <= new_position_y < columns and (new_position_x, new_position_y) not in visited:
                if maze[new_position_x][new_position_y] in (' ', '*'):
                    # Append neighbor node with current node as parent
                    stack.append(Node(new_position_x, new_position_y, maze[new_position_x][new_position_y], node))
    print_path(maze)
    #write_path_to_file(maze)
    print("No more goals were found.")
    return maze
depth_first_search("./Complexity_and_Uninformed_Search_files/data/map1.txt")
#write_to_file("DFS_1", "./Complexity_and_Uninformed_Search_files/data/")
depth_first_search("./Complexity_and_Uninformed_Search_files/data/map2.txt")
#write_to_file("BFS_2", "./Complexity_and_Uninformed_Search_files/data/")
depth_first_search("./Complexity_and_Uninformed_Search_files/data/map3.txt")
#write_to_file("BFS_3", "./Complexity_and_Uninformed_Search_files/data/")
    # YOUR CODE HERE
    #raise NotImplementedError()
```

Output:
```
========================================================================== 

|  s─────────────────────────────│|                                      | 

|                    │────────────|                                      | 

|====================────────────│|                                      | 

|                                ───────────────────────────────────────│| 

|                                                              │─────────| 

|                  ===================                         ─────────│| 

|│───────────────────────────────────────────────────────────────────────| 

|───────────────────────────────────────────────────────────────────────│| 

|│───────────────────────────────────────────────────────────────────────| 

|───────────────│                     ==========================         | 

|               ────────────────────────────────────────────────────────│| 

|│───────────────────────────────────────────────│───────────────────────| 

```

<!-- grade_id: test_uninformed_search_D0, points: 20, grade: true, solution: false -->
```python
# Test

# Please write your code test here!

```

## Task 2.5

**[20 Point(s)]**

## Iterative Deepening Search [20 Points]

Implement *Iterative Deepening Search*. Make sure to fill the docstring of the function and add typehints such that one can understand what arguments are expected and what it returns.

Run the search algorithm for each map and save the results into `./data/results/<search_algo_name>_<map_number>.txt`.

<!-- grade_id: uninformed_search_E, points: 0, grade: false, solution: true -->
```python
def iterative_deepening_depth_first_search(maze_map, monitor=False):
    """ This function implements the iterative deepening depth-first search (IDDFS) algorithm.
    THe key point about IDDFS is that it uses depth-limited depth-first search with an incrementing depth
    limit and thereby explores the node basically breadth first, but at the same time makes use of the space
    saving procedure of DFS, which keeps only b number of nodes in the memory, with b being the branching
    factor of the current node.
    
    Parameters
    ----------
    maze_map : [type]
        [description]
    monitor : bool
        Flag for enabling or disabling command line output for the search run. Defaults to False: No
        monitoring.
    
    Returns
    -------
    [type]
        [description]
    """

    
    

    # YOUR CODE HERE
    #raise NotImplementedError()
```

<!-- grade_id: test_uninformed_search_E0, points: 20, grade: true, solution: false -->
```python
# Test

# Please write your code test here!

```